---
title_meta: 'Chapter 7'
title: Course work
description: 'In this page, you can do your own analysis for the course work.'
---

## Course work

```yaml
type: NormalExercise
key: exercise7_1
lang: r
xp: 0
skills: 1
```
In this page you can do your analysis for the course work. Data sets cardiac and diabetes are available by typing 'cardiac' or 'diabetes'. 

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
cardiac<-read.table("https://raw.githubusercontent.com/paulabergman/intro-to-biostatistics-with-r/master/Cardiac.txt",header=T,row.names=1)
diabetes<-read.table("https://raw.githubusercontent.com/paulabergman/intro-to-biostatistics-with-r/master/Diabetes.txt",header=TRUE,row.names=1)
```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
